# Tenta 2022-06-10

## Part 1

The British pound (GBP) is the oldest existing currency today. Originally the value of a pound note was defined to be equivalent to the value of a pound (in weight) of silver. 

Until the 70's the division of the pound into smaller unit was different from most other currencies: two smaller units of currency was defined as

    1 pound = 20 shillings
    1 shilling = 12 pennies
    
**1** Write functions that handle GBP, assuming that any wallet can be represented by a triplet of ints (pounds, shilling, pence)



<!-- grade_id: cell-2fe171bc6a61c60a, points: 0, grade: false, solution: true -->
```python
def to_pence(pounds, shillings, pence):
    """
    Convert old British currency to smallest unit (p)
    Example

    >>> to_pence(0, 2, 0)
    24
    """
```

<!-- grade_id: cell-40f84e96e1c7af1a, points: 1, grade: true, solution: false -->
```python
from solutions import to_pence
assert to_pence(1, 0, 0) == 240, f'{to_p(1, 0, 0)} != 200'
assert to_pence(0, 1, 0) == 12
assert to_pence(0, 0, 5) == 5
assert to_pence(1, 1, 1) == 253
```

```{admonition} Solution
:class: dropdown
~~~
def to_pence(pounds, shillings, pence):
    """
    Convert old British currency to smallest unit
    
    >>> to_pence(0, 2, 0)
    24
    """
    return (pounds*20 + shillings)*12 + pence
~~~
```

**2** Given any combination of coins, let a reduce function minimize the combination of coins so that the total amount remains constant.

<!-- grade_id: cell-34a877fddba69899, points: 0, grade: false, solution: true -->
```python
def reduce(pounds, shillings, pence):
    """
    Change the moneys first to maximize the number of pound notes,
    then shillings, such that the total amount is constant
    """
```

<!-- grade_id: cell-2bd955e6c650e552, points: 1, grade: true, solution: false -->
```python
from solutions import reduce
assert reduce(0, 0, 13) == (0, 1, 1), f'{reduce(0, 0, 13)} != (0, 1, 1)'
assert reduce(0, 0, 245) == (1, 0, 5), f'{reduce(0, 0, 245)} != (1, 0, 5)'
assert reduce(0, 0, 252) == (1, 1, 0), f'{reduce(0, 0, 252)} != (1, 1, 0)'
```

```{admonition} Solution
:class: dropdown
~~~
def reduce(pounds, shillings, pence):
    """
    Change the moneys first to maximize the number of pound notes,
    then shillings, such that the total amount is constant
    """
    total_pence = to_pence(pounds, shillings, pence)
    # The number of shillings you get out of these pennies
    total_shillings = total_pence // 12
    # The remaining pennies after chaning into shillings
    remaining_pence = total_pence % 12

    # The number of pounds you get out of these shillings
    total_pounds = total_shillings // 20
    # The remaining shillings you get after changing int pound
    remaining_shillings = total_shillings % 20

    return (total_pounds, remaining_shillings, remaining_pence)
`

~~~
```

**3**. What is the maximum number of pound notes that can be change for a given GBP tuple? Define a function to_pounds for this. Ignore the remaining change

<!-- grade_id: cell-9f1761359c8f1937, points: 0, grade: false, solution: true -->
```python
def to_pounds(pounds, shillings, pence):
    """
    Return the max number of pound notes that can be obtained
    """
```

<!-- grade_id: cell-516921e35b32db2b, points: 1, grade: true, solution: false -->
```python
from solutions import to_pounds
assert to_pounds(0, 40, 0) == 2
assert to_pounds(0, 0, 241) == 1
```

```{admonition} Solution
:class: dropdown
~~~
def to_pounds(pounds, shillings, pence):
    """
    Return the max number of pound notes that can be obtained
    """
    return to_pence(pounds, shillings, pence) // 240
~~~
```

**4**. Same as **3** but for shilling coins

<!-- grade_id: cell-ee793cfcadfc0e9d, points: 0, grade: false, solution: true -->
```python
def to_shillings(pounds, shillings, pence):
    """
    Return the max number of pound notes that can be obtained
    """    
    return to_pence(pounds, shillings, pence) // 12
```

<!-- grade_id: cell-818df9f0c638aa16, points: 1, grade: true, solution: false -->
```python
from solutions import to_shillings
assert to_shillings(1, 0, 0) == 20
assert to_shillings(2, 10, 0) == 50
```

```{admonition} Solution
:class: dropdown
~~~
def to_shillings(pounds, shillings, pence):
    """
    Return the max number of pound notes that can be obtained
    """    
    return to_pence(pounds, shillings, pence) // 12
~~~
```

## Part 2

**5** Make a pound class that allows for addition of GBP amounts. Use the tuple in **1** as an instance attribute `data`

<!-- grade_id: cell-9310bacd23ea94c7, points: 0, grade: false, solution: true -->
```python
class GBP:
    ...
```

<!-- grade_id: cell-d50965666a2b7cc5, points: 1, grade: true, solution: false -->
```python
from solutions import GBP
assert GBP(1, 2, 3).data == (1, 2, 3)
```

```{admonition} Solution
:class: dropdown
The intializer must save the input coins as a tuple in an instance attribute `data`
~~~
    def __init__(self, pounds, shillings, pence):
        # Save notes/coins as a tuple "data"
        self.data = (pounds, shillings, pence)
~~~
```

**6** Define addition for this class

<!-- grade_id: cell-e8330777d9b7974f, points: 1, grade: true, solution: false -->
```python
w1 = GBP(1, 0, 0)
w2 = GBP(2, 0, 0)
w3 = w1 + w2
assert w3.data == (3, 0, 0)
```

```{admonition} Solution
:class: dropdown
As first draft that we may add the pounds, shillings, and pence separately
and create a new GBP object with these sum amounts
~~~
    def __add__(self, other):
        pounds = self.data[0] + other.data[0]
        shillings = self.data[1] + other.data[1]
        pence = self.data[2] + other.data[2]
        return GBP(pounds, shillings, pence)
~~~
```

**7** Reuse the reduce function so that the result of a summation minimizes the number of coins

<!-- grade_id: cell-a782bac560bf50a7, points: 1, grade: true, solution: false -->
```python
w1 = GBP(0, 10, 0)
w2 = GBP(0, 10, 0)
w3 = w1 + w2
assert w3.data == (1, 0, 0), f'{w3.data} != (1, 0, 0)'
```

```{admonition} Solution
:class: dropdown
~~~
    def __add__(self, other):
        pounds = self.data[0] + other.data[0]
        shillings = self.data[1] + other.data[1]
        pence = self.data[2] + other.data[2]

        pounds, shillings, pence = reduce(pounds, shillings, pence)
        return GBP(pounds, shillings, pence)
~~~
```

**8**. Define a special method for equality so that two GBP wallets are considered equal if the total amount is the same

<!-- grade_id: cell-8bbdd29edf430714, points: 1, grade: true, solution: false -->
```python
assert GBP(1, 0, 0) == GBP(0, 20, 0)
assert GBP(0, 1, 0) == GBP(0, 0, 12)
```

```{admonition} Solution
:class: dropdown
~~~
    def __eq__(self, other):
        breakpoint()
        return (to_pence(*self.data) == to_pence(*other.data))
~~~
```

## Part 3

**9**. How can the class be updated so that the following unpacking of notes and coins is possible

<!-- grade_id: cell-9446da304b119776, points: 1, grade: true, solution: false -->
```python
pounds, shillings, pence = GBP(1, 2, 3)
assert pounds == 1
assert shillings == 2
assert pence == 3
```

```{admonition} Solution
:class: dropdown
~~~
    def __iter__(self):
        return iter(self.data)
~~~
```

**10**. Assume that we have no tolerance for debt, and that the total amount always of a GBP object always has to be positive. We can define an exception if an amount is negative

~~~
class NegativePounds(Exception):
    pass
~~~

Make sure that it is raised when an amount is negative, example

    >>> GPB(-1, 0, 0)
    ---------------------------------------------------------------------------
    NegativePounds                            Traceback (most recent call last)
    Input In [128], in <cell line: 1>()
    ----> 1 raise NegativePounds()

    NegativePounds: 


<!-- grade_id: cell-208656121e241b1f, points: 1, grade: true, solution: false -->
```python
from solutions import NegativePounds
assert GBP(1, -19, 0) == GBP(0, 1, 0) # This is fine

import pytest
with pytest.raises(NegativePounds):
    GBP(1, -21, 0) # Not fine
```

```{admonition} Solution
:class: dropdown
~~~
    def __init__(self, pounds, shillings, pence):
        if to_pence(pounds, shillings, pence) < 0:
            raise NegativePounds
        ...
~~~
```

**11**. Update the class so that sorting of wallets is possible. Hint: look at the error message you may initially get

<!-- grade_id: cell-73dadfe68c51c42c, points: 1, grade: true, solution: false -->
```python
w1 = GBP(1, 0, 0)
w2 = GBP(1, 0, 1)

assert sorted([w2, w1]) == [w1, w2]
```

```{admonition} Solution
:class: dropdown
The error you get will contain a messege that ends with
~~~
TypeError: '<' not supported between instances 
~~~
So sorted method requires that ordering can be defined through the `<` operator.
We implement that with the `__lt__` method
~~~
    def __lt__(self, other):
        """
        Special method that implements smaller than

        >>> self < other
        True #or False
        """
        return to_pence(*self.data) < to_pence(*other.data)

~~~
```

**12**. Implement also the following inequality

<!-- grade_id: cell-25ff2540ec6badbe, points: 1, grade: true, solution: false -->
```python
assert w1 <= w2
```

```{admonition}
:class: dropdown
Lower than or equal (<=) is implemented by the `__le__` method
~~~
    def __le__(self, other):
        return to_pence(*self.data) < to_pence(*other.data)

~~~
```