# Tenta 2020-06-05

## Part 1

### 1. Word count

Write a function that counts words in text and returns a dictionary which maps the word to the count number (int) of that word

* It should convert words to lower case
* Words should not have newlines
* Words should not contain separation characters like ", . ? !"

<!-- grade_id: cell-86569f3e86491af5, points: 0, grade: false, solution: true -->
```python
import collections

def word_count(text):
    """
    Calculates words in text, return type: dict
    
    >>> word_count('bye bye')
    {'bye': 2}
    >>> word_count('She loves me yeah yeah yeah')
    {'she': 1, 'loves': 1, 'me': 1, 'yeah': 3}
    """
    count = {}
    text = text.lower()
    for punctuation in ',.?!':
        text = text.replace(punctuation, ' ')
        
    for word in text.split():
        count[word] = count.get(word, 0) + 1
    return count

# try it out
```

Output:
```
{'hello': 1, 'there': 1, 'world': 1}
```

```{dropdown} Here's my dropdown
~~~
And here's my dropdown content
~~~
```

```python
print('foo')
```

Output:
```
foo
```

<!-- grade_id: cell-2ffd84f38aa2115c, points: 1, grade: true, solution: false -->
```python
assert word_count('bye bye') == {'bye': 2}
```

<!-- grade_id: cell-d50cc51cf6cacb8f, points: 1, grade: true, solution: false -->


### 2. Pirate language translator
 
Double any consonant and put an o in between

~~~
>>> to_pirate('hello')
hohelollolo
~~~

The output should preserve leading upper case of the characters in the beginning of the sentance. See the tests below

<!-- grade_id: cell-b625edd3a750a4a0, points: 0, grade: false, solution: true -->
```python
VOWELS = 'eiyaou'
def to_pirate(text):
    pirate = ""
    for char in text:
        # do not duplicae vowels and spaces
        if char not in VOWELS + " ":
            pirate += char + 'o' + char.lower()
        else:
            pirate += char
    return pirate
    
to_pirate('Hello there')    
```

Output:
```
'Hohelollolo tothoherore'
```

<!-- grade_id: cell-d67fcfee6f28fbf1, points: 1, grade: true, solution: false -->
```python
assert to_pirate('hello') == 'hohelollolo'
```

<!-- grade_id: cell-9c71d53ef7a303c4, points: 1, grade: true, solution: false -->
```python
assert to_pirate('Hello there') == 'Hohelollolo tothoherore'
```

## Part 2

### 1. Classes
Define a class Circle that stores position in  the x-y plane and its radius

<!-- grade_id: cell-c9942f7528c934e8, points: 0, grade: false, solution: true -->
```python

from pytest import approx
from math import pi

class Circle:
    # YOUR CODE HERE
```

Let the deault behaviour be a unit circle centered around origo

<!-- grade_id: cell-ff681a9497b11cfd, points: 1, grade: true, solution: false -->
```python
circ = Circle()
assert circ.x == approx(0.0)
assert circ.y == approx(0.0)
assert circ.r == approx(1.0)
```

When we supply arguments to the class constructor the circle shape and position is saved

<!-- grade_id: cell-b162df80c36752d8, points: 1, grade: true, solution: false -->
```python
circ = Circle(r=2)
assert circ.x == approx(0.0)
assert circ.y == approx(0.0)
assert circ.r == approx(2.0)

circ = Circle(x=1, y=1)
assert circ.x == approx(1.0)
assert circ.y == approx(1.0)
assert circ.r == approx(1.0)
```

Add class methods for calculating the area and circumference

<!-- grade_id: cell-a4bf5c652a30a608, points: 1, grade: true, solution: false -->
```python
assert Circle().area() == approx(pi)
assert Circle().circumference() == approx(2*pi)
```

Define an overloading method such that we can compare two Circle objects, being equal if they have same radius and center. Use any comparison between floats with the pytest approx function

<!-- grade_id: cell-f1f8d34fc5da1de6, points: 1, grade: true, solution: false -->
```python
circ_a = Circle(x=1, y=2, r=3)
circ_b = Circle(x=1, y=2, r=3)
assert circ_a == circ_b
```

## Part 3

### 1. Decorator
Consider a mathematical function of a variable $f(x)$. The derivative of the function is a measure of change, illustrated by the slope of a tangent line at a point on the curve

For example given

$$ f(x) = x^2 $$

the derivative is

$$ f'(x) = 2x $$

which means that at e.g. $x_0=3$ a tangent to the curve has slope $f'(3) = 2*3=6$ and has the equation for a straight line

$$ y = f(x_0) + f'(x_0)(x-x_0)$$

```python
import math

import matplotlib.pyplot as plt
import numpy as np
```

```python
def f(x):
    return x**2

def fp(x):
    return 2*x

x = np.arange(0, 5, .1)
x0 = 3.0

plt.plot(x, f(x))
plt.plot(x, f(x0) + fp(x0)*(x-x0))
```

Mathematically derivatives are understood as the ratio of differences in y and x values close to a given point and taking the limit where the point get closer

$$ f'(x) \approx \frac{f(x + \delta) - f(x-\delta)}{2\delta} $$

For a small number $\delta$ this can be used as an approximation to the derivate for a function

```python
delta = 0.1

plt.vlines(x0-delta, 0, f(x0-delta), linestyle='dashed')
plt.vlines(x0+delta, 0, f(x0+delta), linestyle='dashed')
plt.hlines(f(x0-delta), 0, x0-delta, linestyle='dashed')
plt.hlines(f(x0+delta), 0, x0+delta, linestyle='dashed')

plt.text(2.9, -1, '$2\delta$')
plt.text(0, 10, '$f(3 + \delta)$')
plt.text(0, 7, '$f(3 -\delta)$')


plt.plot(x, f(x))
#plt.plot(dx, f(dx))
```

### Task

Write a decorator that implements the numerical approximattion to a derivative

The decorator applied to a library function, e.g. math.sin should then return a function for which approximately

    derivative(math.sin)(x) ≈ math.cos(x)

for most values of $x$

Likewise the decorator applied to a function definition with the @-syntax

~~~
@derivative
def f(x):
     return x**2
~~~

defines a function that satisfies

    f(x) ≈ 2*x
    
for most values of $x$

<!-- grade_id: cell-42bff3a7cd29f897, points: 0, grade: false, solution: true -->
```python
delta = 1e-8

# YOUR CODE HERE
```

<!-- grade_id: cell-656f185b00847d9c, points: 1, grade: true, solution: false -->
```python
epsilon = 1e-6

@derivative
def f(x):
    return x**2

assert abs(f(3.0) - 6.0) < epsilon
```

<!-- grade_id: cell-656f185b00847d9d, points: 1, grade: true, solution: false -->
```python
assert abs(derivative(math.sin)(.5) - math.cos(.5)) < epsilon
```

### 2. Generator


This is a filter that returns lines of a file that have a length greader than maximum

If the content of a textfile `file.txt` is

~~~
hello 
you
merry
~~~

We except a function `find_long_lines` to behave like this

~~~
>>> for line in find_long_lines('file.txt', 3):
...    print(line)
hello
merry
~~~

Knowing that the return value of the function must be iterable to use in a for-loop. 
Implement this function as a generator.

The max_length should be without the final newline character `\n` of the line 

<!-- grade_id: cell-40e42ac0ba0a1d0e, points: 0, grade: false, solution: true -->
```python
def find_long_lines(filename, max_length):
    # YOUR CODE HERE
```

```python
%%file file.txt
hello
you
merry
```

<!-- grade_id: cell-f5b0233fb579e1ad, points: 1, grade: true, solution: false -->
```python
assert list(find_long_lines('file.txt', 3)) == ['hello', 'merry']

```

<!-- grade_id: cell-87b704bdf07423c3, points: 1, grade: true, solution: false -->
```python

```