Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
```

---

# Exam in BB1000, 2021-06-01

## Grading

There are three levels in order of increasing difficulty

* Grade D for 3 of 4 points at level 1 and 1 point at level 2
* Grade C for 3 of 4 points at levels 1 and 2
* Grade B as C plus 1 point at level 3
* Grade A for 3 of 4 points at all levels

# Requirements

For the whole exam you should be able to import the following libraries

```python
import numpy
import requests
import IPython
```

# Level 1


**1. Given that we have a list of email addresses in a file, in this example with the name 'email.txt'**

```python
%%file emails.txt
bonnie@kth.se
clyde@chalmers.se
scarlett@kth.se
rhett@chalmers.se
```

Output:
```
Overwriting emails.txt
```

**write a function that extracts the account names (the part to the left of the '@' sign) in a list**

~~~
>>> get_usernames('emails.txt')
['bonnie', 'clyde']
~~~

<!-- grade_id: cell-f822b8fd33ae224e, points: 0, grade: false, solution: true -->
```python
def get_usernames(source: str) -> list:
    """
    Returns username from a file of mail addresses as a list
    """
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-6d536596d441b56e, points: 1, grade: true, solution: false -->
```python
users = get_usernames("emails.txt")
expected = ["bonnie", "clyde", "scarlett", "rhett"]
assert users == expected, f'\n{users} \n!= \n{expected}'

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-4-8e4c6e1dcfe5> in <module>
      1 users = get_usernames("emails.txt")
      2 expected = ["bonnie", "clyde", "scarlett", "rhett"]
----> 3 assert users == expected, f'\n{users} \n!= \n{expected}'

AssertionError: 
None 
!= 
['bonnie', 'clyde', 'scarlett', 'rhett']
```

<!-- grade_id: cell-6ac45588bcc8080d, points: 1, grade: true, solution: false -->
```python

```

**2. Group the users by affiliation with a function that returns a dictionary**

~~~
>>> get usernames_by_affiliation('emails.txt')
{'kth': ['bonnie', 'scarlett'], 'chalmers': ['clyde', 'rhett']}
~~~

<!-- grade_id: cell-252f3ca6ac3a00ae, points: 0, grade: false, solution: true -->
```python
def get_usernames_by_affiliation(source: str) -> list:
    """
    Returns usernames from a file grouped by affiliation
    """
    # YOUR CODE HERE
    

```

<!-- grade_id: cell-4485fa612388c633, points: 1, grade: true, solution: false -->
```python
users = get_usernames_by_affiliation('emails.txt')
expected = {'kth': ['bonnie', 'scarlett'], 'chalmers': ['clyde', 'rhett']}  
assert users == expected, f'\n{users} \n!= \n{expected}'
```

<!-- grade_id: cell-252f3ca6ac3a00af, points: 1, grade: true, solution: false -->
```python

```

# Level 2

## Classes


Based on dictionary input in the following form

~~~
{
    'givenName': 'Sigbritt',
    'familyName': 'Karlsson',
    'url': 'https://www.kth.se/profile/sigbritt',
    'email': 'rektor@kth.se',
}
~~~

**1. Create an Employee class that takes this type of record data as input and which stores the values in instance attributes `first`, `last`, `email`**

<!-- grade_id: cell-5bed954dde7d9e1e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-86051b003641f730, points: 1, grade: true, solution: false -->
```python
record = {
    'givenName': 'Sigbritt',
    'familyName': 'Karlsson',
    'url': 'https://www.kth.se/profile/sigbritt',
    'email': 'rektor@kth.se',
}

args = (record['givenName'], record['familyName'], record['email'])
rector = Employee(*args)
assert rector.first == 'Sigbritt'
assert rector.last == 'Karlsson'
assert rector.email == 'rektor@kth.se'
```

**2. Implement instance methods such that the following relations hold**

<!-- grade_id: cell-3cb30633bae04d4f, points: 1, grade: true, solution: false -->
```python
emp = Employee('Malvina', 'Sievert', 'malvina@kth.se')
assert str(emp) == 'Malvina Sievert <malvina@kth.se>', f'"{str(emp)}" is incorrect'
assert repr(emp) == 'Employee("Malvina", "Sievert", "malvina@kth.se")', f'"{repr(emp)}" is incorrect'
```

**3. Have a class variable `all` be a list which contain references to all created Employee records**

<!-- grade_id: cell-96206ee5afd04747, points: 1, grade: true, solution: false -->
```python
Employee.all.clear()
assert len(Employee.all) == 0, f'{len(Employee.all)} != 0'
Employee(*args)
assert len(Employee.all) == 1, f'{len(Employee.all)} != 1'
```

**4. A method `to_mail_file` that takes the current state of the Employee class and writes a maillist file (as in ex 1)**

<!-- grade_id: cell-2e3e6c6046d9cc36, points: 1, grade: true, solution: false -->
```python
Employee.all.clear()
Employee('Chip', '', 'chip@kth.se')
Employee('Dale', '', 'dale@kth.se')

Employee.to_mail_file('employees.txt')

assert get_usernames('employees.txt') == ['chip', 'dale'], f'Incorrect: {get_usernames("employees.txt")}'
```

# Level 3

A decorator is normally applied in the file where a function is defined with the `@decorator` syntax on the line above the function definition. 

~~~
@decorator
def f():
    ...
~~~

which is equivalent to a redefinition

```
f = decorator(f)
```

**1. Write a `ChalmersHack` decorator that modifiees all email addresses to
end with `@chalmers.se`**



<!-- grade_id: cell-d56d07f0f8902c20, points: 0, grade: false, solution: true -->
```python
def ChalmersHack(cls):
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-d40ea7172d8e4cdb, points: 1, grade: true, solution: false -->
```python
_Employee = ChalmersHack(Employee)
emp = _Employee('John', 'Doe', 'jd@kth.se')
assert emp.email == 'jd@chalmers.se'
```

**2. Write a function that looks up usernames in the KTH public database**

Example:

~~~
>>> requests.get('https://api.kth.se/api/profile/1.1/sigbritt').json()
{'givenName': 'Sigbritt',
 'familyName': 'Karlsson',
 'url': 'https://www.kth.se/profile/sigbritt',
 'email': 'rektor@kth.se',
 'image': 'https://www.kth.se/files/avatar/sigbritt',
 'telephone': '087907001',
 'jobTitle': 'PROFESSOR,REKTOR',
 'jobTitle-en': 'PRESIDENT,PROFESSOR',
 'worksFor': [{'url': 'https://www.kth.se/directory/t/ta',
   'name': 'UNIVERSITETSLEDNING',
   'name-en': 'UNIVERSITETSLEDNING'}],
 'workLocation': 'BRINELLVÄGEN 8; TEKNIKRINGEN 56-58'}
~~~

First verify that the output in the cell below looks the same

```python
requests.get('https://api.kth.se/api/profile/1.1/sigbritt').json()
```

**3. Write a generator functinon that takes a sequence of usernames as inputs that yields Employee record and the image urls as output**

<!-- grade_id: cell-d6ff720ac4b319c9, points: 0, grade: false, solution: true -->
```python
def get_users(users):
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-58810442437f174d, points: 1, grade: true, solution: false -->
```python
users = get_users(['sigbritt', 'mostling'])
result = next(users)
expected = Employee('Sigbritt', 'Karlsson', 'rektor@kth.se'), 'https://www.kth.se/files/avatar/sigbritt'

assert result[1] == expected[1], f'\n{result[1]} \n!=\n{expected[1]}'
```

<!-- grade_id: cell-a10526585a12624e, points: 1, grade: true, solution: false -->
```python

```

Verify that you get images of the heads of KTH Schools with the final code 

```python
skolchefer = ['muriel', 'mil', 'berlijn', 'parj', 'dirocco']

for record, image in get_users(skolchefer):
    display(IPython.display.Image(url=image))
    print(record.first, record.last)

```

<!-- grade_id: cell-a873ec5d8f44004b, points: 1, grade: true, solution: false -->
```python
assert len(list(get_users(skolchefer))) == 5
```