Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
PERSON_NUMBER = ""
```

---

## Tenta BB1000 2022-08-19

### Grading

There are three levels in order of increasing difficulty

    Grade E for 75% at level 1
    Grade D as E plus 25% at level 2
    Grade C for 75% at levels 1 and 2
    Grade B as C 25% at level 3
    Grade A for 75% at all levels



### Part 1

a) Define a `minmax` function that returns a tuple containing the smallest and largest number of a sequence such that

    >>> minmax([1, 3, 2, 5])
    (1, 5)

<!-- grade_id: cell-d8921dbcde639869, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
################
```

<!-- grade_id: cell-3bcdfa3488f30b99, points: 1, grade: true, solution: false -->
```python
result = minmax([1, 3, 2, 5])
expected = (1, 5)
assert result == expected, f'\nresult  : {result}\nexpected: {expected}'

result = minmax((9, 7, 6))
expected = (6, 9)
assert result == expected, f'\nresult  : {result}\nexpected: {expected}'

```

b) Define an optional parameter `default` to be used in the return value if the sequence is empty. If the value is not given the min and max values are set to `None` in this case

    >>> minmax([1, 2])
    (1, 2)
    >>> minmax([])
    (None, None)
    >>> minmax([], default=0)
    (0, 0)

<!-- grade_id: cell-77a2d945f091a761, points: 1, grade: true, solution: false -->
```python
result = minmax([])
expected = (None, None)
assert result == expected, f'\nresult  : {result}\nexpected: {expected}'

result = minmax([], default=0)
expected = (0, 0)
assert result == expected, f'\nresult  : {result}\nexpected: {expected}'

```

c) Make sure that the function hold for strings, reporting first and last words put in alphabetical order

    >>> minmax(['world', 'is', 'weird', 'hello'])
    ('hello', 'world')

<!-- grade_id: cell-37939ec80886d32c, points: 1, grade: true, solution: false -->
```python
result = minmax(['world', 'is', 'strange', 'hello'])
expected = ('hello', 'world')

assert result == expected, f'\nresult  : {result}\nexpected: {expected}'
```

d) Define an optional keyword `key` for a callable that defines sorting order, e.g. when sorting by word length we may write

    >>> minmax(['world', 'is', 'strange', 'hello'], key=len)
    ('is', 'strange')
    
and let its default value be `None`

<!-- grade_id: cell-02c3f52b78b218ec, points: 1, grade: true, solution: false -->
```python
result = minmax(['world', 'is', 'strange', 'hello'], key=len)
expected = ('is', 'strange')

assert result == expected, f'\nresult  : {result}\nexpected: {expected}'

result = minmax(['World', 'is', 'strange', 'hello'], key=str.lower)
expected = ('hello', 'World')

assert result == expected, f'\nresult  : {result}\nexpected: {expected}'
```

### Part 2

Build on the previous. This time we consider special functions for dictionaries but that call your `minmax` under the hood

    >>> minmax_keys({'a': 3, 'b': 4, 'c': 5})
    ('a', 'c')
    >>> minmax_values{'a': 3, 'b': 4, 'c': 5})
    (3, 5)
    >>> keys_of_minmax_values{'a': 2, 'b': 1})
    ('b', 'a')
    >>> minmax_items({'b': 2, 'a': 1})
    [('a', 1), ('b', 2)]



a) Let `minmax_keys` be a funtion that returns the min and max keys of a dictionary

<!-- grade_id: cell-0e953b20316873ee, points: 0, grade: false, solution: true -->
```python
def minmax_keys(dct):
    # YOUR CODE HERE
    ################
```

<!-- grade_id: cell-cadac0ae5c3d0945, points: 1, grade: true, solution: false -->
```python
assert minmax_keys(dict(a=1, b=2)) == ('a', 'b')
assert minmax_keys({'a': 3, 'b': 4, 'c': 5}) == ('a', 'c')

# this verifies that minmax was called by minmax_keys
from unittest.mock import patch
with patch('__main__.minmax', return_value=(None, None)) as mock_mm:
    minmax_keys({'a': 1})
    
mock_mm.assert_called_with({'a': 1})
```

b) Let `minmax_values` be a function that returns the min and max values of a dictionary

<!-- grade_id: cell-368fa25ee36fb54a, points: 0, grade: false, solution: true -->
```python
def minmax_values(dct):
    # YOUR CODE HERE
    ################
```

<!-- grade_id: cell-ecd6cb4b9edd0928, points: 1, grade: true, solution: false -->
```python
assert minmax_values(dict(a=1, b=2)) == (1, 2)
assert minmax_values({'a': 3, 'b': 4, 'c': 5}) == (3, 5)

# did you call minmax?

with patch('__main__.minmax', return_value=(None, None)) as mock_mm:
    minmax_values({})
    
mock_mm.assert_called()
```

c) Let `keys_of_minmax_values` be a function that returns the keys of the smallest and largest values in the dictionary. Hint: you will need to define a helper function to pass to `minmax`

<!-- grade_id: cell-100fb43e7ece47e8, points: 0, grade: false, solution: true -->
```python
def keys_of_minmax_values(dct):
    # YOUR CODE HERE
    ################

```

<!-- grade_id: cell-456dbc3c0a012bed, points: 1, grade: true, solution: false -->
```python
dct = {'a': 2, 'b': 1}
assert minmax(dct) == ('a', 'b')
assert minmax_values(dct) == (1, 2)
assert keys_of_minmax_values(dct) == ('b', 'a')

# did you call minmax?

with patch('__main__.minmax', return_value=(None, None)) as mock_mm:
    keys_of_minmax_values({})
    
mock_mm.assert_called()
```

d) A `minmax_items` will return key-value pairs of a dictionary corresponding to the smallest and largest keys respectively

<!-- grade_id: cell-2583e60042563a75, points: 0, grade: false, solution: true -->
```python
def minmax_items(dct):
    # YOUR CODE HERE
    ################

```

<!-- grade_id: cell-387944423ac9aba7, points: 1, grade: true, solution: false -->
```python
dct = {'b': 2, 'a': 1}
assert minmax_items(dct) == (('a', 1), ('b', 2))

# did you call minmax?

with patch('__main__.minmax', return_value=(None, None)) as mock_mm:
    minmax_items({})
    
mock_mm.assert_called()
```

### Part 3

Build on the previous example again. Consider the official documentation for [named tuples](https://docs.python.org/3/library/collections.html#collections.namedtuple)

The following function `minmax2` calls the function in part 1, works the same way as before, but returns an object with min, and max attributes

    >>> minmax2([1, 2, 3]).min
    1
    >>> minmax2([1, 2, 3]).max
    2
    
a) A new function `minmax2` with same arguments satisfies the same tests but that is returning a named tuple (define it as `MinMax`)

<!-- grade_id: cell-fe5adbbd6d156f14, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
################
```

<!-- grade_id: cell-760913208c76b573, points: 1, grade: true, solution: false -->
```python
# repeat the previous test, 
result = minmax2(['world', 'is', 'strange', 'hello'], key=len)
expected = ('is', 'strange')

assert result == expected, f'\nresult  : {result}\nexpected: {expected}'

# but note that the return value is not the ordinary tuple
assert result.__class__.__name__ == 'MinMax'
```

b) The return value now has attributes `min` and `max`

<!-- grade_id: cell-aef9482444ec1821, points: 1, grade: true, solution: false -->
```python
# the return value of the `minmax2` function is an object with attributes min and max

result = minmax2(['world', 'is', 'strange', 'hello'], key=len)

assert result.min == 'is'
assert result.max == 'strange'


```

c) Was the original function called? 

<!-- grade_id: cell-7591cc9cb3b5f76b, points: 1, grade: true, solution: false -->
```python
# This test verifies that the original function minmax was called

from unittest.mock import patch
with patch('__main__.minmax', return_value=(None, None)) as mock_mm:
    minmax2([])
    
mock_mm.assert_called_with([], default=None, key=None)
```

d) let minmax2 take any number of positional arguments. If there is one argument it is assumed to be an iterable. If there is more than one arguments, the items are compared to each other

<!-- grade_id: cell-0a4275a5867cc18b, points: 1, grade: true, solution: false -->
```python
assert minmax2([1, 2, 3]) == (1, 3)
assert minmax2(1, 2, 3) == (1, 3)

assert minmax2(range(5)) == (0, 4)
assert minmax2(*range(5)) == (0, 4)
```