# Tenta 2021-06-01
*Personal information modified from the original*

## Requirements

For the whole exam you should be able to import the following libraries

```python
import numpy
import requests
import IPython
```

## Level 1


**1. Given that we have a list of email addresses in a file, in this example with the name 'email.txt'**

```python
%%file emails.txt
bonnie@kth.se
clyde@chalmers.se
scarlett@kth.se
rhett@chalmers.se
```

Output:
```
Overwriting emails.txt
```

**write a function that extracts the account names (the part to the left of the '@' sign) in a list**

~~~
>>> get_usernames('emails.txt')
['bonnie', 'clyde']
~~~

<!-- grade_id: cell-f822b8fd33ae224e, points: 0, grade: false, solution: true -->
```python
def get_usernames(source: str) -> list:
    """
    Returns username from a file of mail addresses as a list
    """
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-6d536596d441b56e, points: 1, grade: true, solution: false -->
```python
from solutions import get_usernames
users = get_usernames("emails.txt")
expected = ["bonnie", "clyde", "scarlett", "rhett"]
assert users == expected, f'\n{users} \n!= \n{expected}'

```

<!-- grade_id: cell-6ac45588bcc8080d, points: 1, grade: true, solution: false -->
```python

```

```{admonition} Solution
:class: dropdown
~~~
def get_usernames(source: str) -> list:
    """               
    Returns usernames from a file of mail addresses as a list
    """
    with open(source) as f:                                     
        usernames = []
        for line in f:
            username = line.split('@')[0]
            usernames.append(username)
    return users               
~~~
```

**2. Group the users by affiliation with a function that returns a dictionary**

~~~
>>> get usernames_by_affiliation('emails.txt')
{'kth': ['bonnie', 'scarlett'], 'chalmers': ['clyde', 'rhett']}
~~~

<!-- grade_id: cell-252f3ca6ac3a00ae, points: 0, grade: false, solution: true -->
```python
def get_usernames_by_affiliation(source: str) -> list:
    """
    Returns usernames from a file grouped by affiliation
    """
    # YOUR CODE HERE
    

```

<!-- grade_id: cell-4485fa612388c633, points: 1, grade: true, solution: false -->
```python
from solutions import get_usernames_by_affiliation
users = get_usernames_by_affiliation('emails.txt')
expected = {'kth': ['bonnie', 'scarlett'], 'chalmers': ['clyde', 'rhett']}  
assert users == expected, f'\n{users} \n!= \n{expected}'
```

<!-- grade_id: cell-252f3ca6ac3a00af, points: 1, grade: true, solution: false -->
```python

```

```{admonition} Solution
:class: dropdown
~~~
def get_usernames_by_affiliation(source: str) -> dict:
    """
    Returns usernames from a file as a dictionary grouped by affiliation
    """
    users = {}
    with open(source) as f:
        for line in f:
            username, domain = line.split('@')
            affiliation, _ = domain.split('.')
            if affiliation not in users:
                users[affiliation] = []
            users[affiliation].append(username)
    return users
~~~
```

## Level 2

### Classes


Based on dictionary input in the following form

~~~
{
    'givenName': 'Anders',
    'familyName': 'Söderholm',
    'url': 'https://www.kth.se/profile/soderhol',
    'email': 'rektor@kth.se',
}
~~~

**1. Create an Employee class that takes this type of record data as input and which stores the values in instance attributes `first`, `last`, `email`**

<!-- grade_id: cell-5bed954dde7d9e1e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-86051b003641f730, points: 1, grade: true, solution: false -->
```python
from solutions import Employee
record = {
    'givenName': 'Anders',
    'familyName': 'Söderholm',
    'url': 'https://www.kth.se/profile/soderhol',
    'email': 'rektor@kth.se',
}

args = (record['givenName'], record['familyName'], record['email'])
rector = Employee(*args)
assert rector.first == 'Anders'
assert rector.last == 'Söderholm'
assert rector.email == 'rektor@kth.se'
```

```{admonition} Solution
:class: dropdown
~~~
class Employee:
    def __init__(self, first, last, email):
        self.first = first
        self.last = last
        self.email = email
~~~
```

**2. Implement instance methods such that the following relations hold**

<!-- grade_id: cell-3cb30633bae04d4f, points: 1, grade: true, solution: false -->
```python
emp = Employee('Malvina', 'Sievert', 'malvina@kth.se')
assert str(emp) == 'Malvina Sievert <malvina@kth.se>', f'"{str(emp)}" is incorrect'
assert repr(emp) == 'Employee("Malvina", "Sievert", "malvina@kth.se")', f'"{repr(emp)}" is incorrect'
```

```{admonition} Solution
:class: dropdown
~~~
class Employee
    ...
    def __str__(self):
        return f'{self.first} {self.last} <{self.email}>'

    def __repr__(self):
        return f'Employee("{self.first}", "{self.last}", "{self.email}")'
~~~
```

**3. Have a class variable `all` be a list which contain references to all created Employee records**

<!-- grade_id: cell-96206ee5afd04747, points: 1, grade: true, solution: false -->
```python
Employee.all.clear()
assert len(Employee.all) == 0, f'{len(Employee.all)} != 0'
Employee(*args)
assert len(Employee.all) == 1, f'{len(Employee.all)} != 1'
```

```{admonition} Solution
:class: dropdown
~~~
class Employee:
    all = []
    def __init__(self, first, last, email):
        self.first = first
        self.last = last
        self.email = email
        Employee.all.append(self)
    ...
~~~
```

**4. A method `to_mail_file` that takes the current state of the Employee class and writes a maillist file (as in ex 1)**

<!-- grade_id: cell-2e3e6c6046d9cc36, points: 1, grade: true, solution: false -->
```python
Employee.all.clear()
Employee('Chip', '', 'chip@kth.se')
Employee('Dale', '', 'dale@kth.se')

Employee.to_mail_file('employees.txt')

assert get_usernames('employees.txt') == ['chip', 'dale'], f'Incorrect: {get_usernames("employees.txt")}'

```

```{admonition} Solution
:class: dropdown
~~~
class Employee:
    ...
    @staticmethod
    def to_mail_file(filename):
        with open(filename, 'w') as f:
            for e in Employee.all:
                print(e.email, file=f)
~~~
```

## Level 3

A decorator is normally applied in the file where a function is defined with the `@decorator` syntax on the line above the function definition. 

~~~
@decorator
def f():
    ...
~~~

which is equivalent to a redefinition

```
f = decorator(f)
```

**1. Write a `ChalmersHack` decorator that modifiees all email addresses to
end with `@chalmers.se`**



<!-- grade_id: cell-d56d07f0f8902c20, points: 0, grade: false, solution: true -->
```python
def ChalmersHack(cls):
    # YOUR CODE HERE
    ...
    
```

<!-- grade_id: cell-d40ea7172d8e4cdb, points: 1, grade: true, solution: false -->
```python
from solutions import ChalmersHack
_Employee = ChalmersHack(Employee)
emp = _Employee('John', 'Doe', 'jd@kth.se')
assert emp.email == 'jd@chalmers.se'
```

```{admonition} Solution
:class: dropdown
~~~
def ChalmersHack(cls):
    def wrapper(*args, **kwargs):
        employee = cls(*args, **kwargs)
        employee.email = employee.email.split('@')[0] + '@chalmers.se'
        return employee
    return wrapper
~~~
```

**2. Write a function that looks up usernames in the KTH public database**

Example:

~~~
>>> requests.get('https://api.kth.se/api/profile/1.1/soderholigbritt').json()
{'givenName': 'Anders',
 'familyName': 'Söderholm',
 'url': 'https://www.kth.se/profile/soderhol',
 'email': 'rektor@kth.se',
 'image': 'https://www.kth.se/files/avatar/soderhol',
 'telephone': '087907001',
 'jobTitle': 'REKTOR',
 'jobTitle-en': '',
 'worksFor': [{'url': 'https://www.kth.se/directory/t/ta',
   'name': 'UNIVERSITETSLEDNING',
   'name-en': 'PRESIDENT AND MANAGEMENT'}],
 'workLocation': 'BRINELLVÄGEN 8'}
~~~

First verify that the output in the cell below looks the same

```python
requests.get('https://api.kth.se/api/profile/1.1/soderhol').json()
```

Output:
```
{'givenName': 'Anders',
 'familyName': 'Söderholm',
 'url': 'https://www.kth.se/profile/soderhol',
 'email': 'rektor@kth.se',
 'image': 'https://www.kth.se/files/avatar/soderhol',
 'telephone': '087907001',
 'jobTitle': 'REKTOR',
 'jobTitle-en': '',
 'worksFor': [{'url': 'https://www.kth.se/directory/t/ta',
   'name': 'UNIVERSITETSLEDNING',
   'name-en': 'PRESIDENT AND MANAGEMENT'}],
 'workLocation': 'BRINELLVÄGEN 8'}
```

**3. Write a generator functinon that takes a sequence of usernames as inputs that yields Employee record and the image urls as output**

<!-- grade_id: cell-d6ff720ac4b319c9, points: 0, grade: false, solution: true -->
```python
def get_users(users):
    # YOUR CODE HERE
    ...
    
```

<!-- grade_id: cell-58810442437f174d, points: 1, grade: true, solution: false -->
```python
from solutions import get_users
users = get_users(['soderhol', 'mil'])
result = next(users)
expected = Employee('Anders', 'Söderholm', 'rektor@kth.se'), 'https://www.kth.se/files/avatar/soderhol'

assert result[1] == expected[1], f'\n{result[1]} \n!=\n{expected[1]}'
```

<!-- grade_id: cell-a10526585a12624e, points: 1, grade: true, solution: false -->
```python

```

```{admonition} Solution
:class: dropdown
~~~
def get_users(users):
    url = 'https://api.kth.se/api/profile/1.1/'
    for u in users:
        r = requests.get(url + u).json()
        yield Employee(r['givenName'], r['familyName'], r['email']), r['image']
~~~
```

Verify that you get images of the heads of KTH Schools with the final code 

```python
skolchefer = ['bber', 'ameliek', 'alz', 'parj', 'dirocco']

for record, image in get_users(skolchefer):
    display(IPython.display.Image(url=image))
    print(record.first, record.last)

```

Output:
```
<IPython.core.display.Image object>
Björn Berggren
<IPython.core.display.Image object>
Amelie Eriksson Karlström
<IPython.core.display.Image object>
Ann Lantz
<IPython.core.display.Image object>
Pär Jönsson
<IPython.core.display.Image object>
Sandra Di Rocco
```

<!-- grade_id: cell-a873ec5d8f44004b, points: 1, grade: true, solution: false -->
```python
assert len(list(get_users(skolchefer))) == 5
```