Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
PERSONNUMBER = ""
```

---

# Tenta 2023-03-17 14:00

To intialize this exam to not forget run the import cell next

```python
import time
```

## Part 1


### A digital clock

The time.time function returns the number of seconds since a reference point of time known as the UNIX epoch, 1 Jan 1970 (GMT).

    >>> time.time()
    1678896639.9824622


Write a function `current_time_of_day` that has no required input, and that returns its output as a string-representation the of current time (GMT). Example:

    >>> current_time_of_day()
    '13:05:00'
    
will be the output 5 minutes into the time of the exam (still GMT).

You are expected to use `time.time` in your solution and to truncate the `float` value it returns to an integer - i.e. remove/ignore decimals.

<!-- grade_id: cell-e2b5da5074216899, points: 0, grade: false, solution: true -->
```python
def current_time_of_day(tz=0):
    """
    >>> current_time_of_day()
    '13:05:00'
    """
    # YOUR CODE HERE
    ################
```

### Verify return type

Here we verify that we have string of length 8 and that the the output at midnight GMT is

    >>> current_time_of_day()
    '00:00:00'

<!-- grade_id: cell-382721a62e0b5d1a, points: 1, grade: true, solution: false -->
```python
assert type(current_time_of_day()) == str and len(current_time_of_day()) == 8


```

### Verify before/after midnight

At any midnight the return value should  be the string

    '00:00:00'
    
and a second before that

    '23:59:59'

<!-- grade_id: cell-124ab4469beaee02, points: 1, grade: true, solution: false -->
```python
import random
from unittest.mock import patch

any_random_midnight = 60*60*24*random.randint(0, 1000)

# This make time.time() return a chosen value during a test

with patch('time.time', return_value=any_random_midnight):  
    calculated = current_time_of_day()
    expected = '00:00:00'
    assert calculated == expected, f'{calculated=} != {expected}'
    
just_before_midnight = any_random_midnight - 1

with patch('time.time', return_value=just_before_midnight): 
    calculated = current_time_of_day()
    expected = '23:59:59'
    assert calculated == expected, f'{calculated=} != {expected=}'
  

```

### Verify handling of fractional seconds

Verify that fractional values from `time.time` is handled such that 

    .9 seconds after midnight -> "00:00:00"
    1.1 seconds after midnight -> "00:00:01"

<!-- grade_id: cell-5a332274f4db5a22, points: 1, grade: true, solution: false -->
```python
import random
from unittest.mock import patch

any_random_midnight = 60*60*24*random.randint(0, 1000)

# This make time.time() return a chosen value during a test

with patch('time.time', return_value=any_random_midnight + 0.9):  
    calculated = current_time_of_day()
    expected = '00:00:00'
    assert calculated == expected, f'{calculated=} != {expected}'
    
with patch('time.time', return_value=any_random_midnight + 1.1):  
    calculated = current_time_of_day()
    expected = '00:00:01'
    assert calculated == expected, f'{calculated=} != {expected}'
```

### Include optional timezone

Complement the function with an optional argument `tz`, an integer representing hours ahead of GMT. For Central European Time (CET) this has value 1. 

The docstring above could be updated with 

    >>> current_time_of_day()
    '13:05:00'
    >>> current_time_of_day(tz=1)
    '14:05:00'


Verify for yourself that it returns the current time in CET timezone

<!-- grade_id: cell-e41c8f291cbbf177, points: 1, grade: true, solution: false -->
```python
import random
from unittest.mock import patch

any_random_midnight = 60*60*24*random.randint(0, 1000)

# This make time.time() return a chosen value during a test

with patch('time.time', return_value=any_random_midnight):  
    calculated = current_time_of_day(tz=1)
    expected = '01:00:00'
    assert calculated == expected, f'{calculated=} != {expected}'
    
with patch('time.time', return_value=any_random_midnight):  
    calculated = current_time_of_day(tz=2)
    expected = '02:00:00'
    assert calculated == expected, f'{calculated=} != {expected}'

```

```python
RNA = 'AUGUCGGAAAAAGAAAUUUGGGAAAAAGUGCUUGAAAUUGCUCAAGAAAAAUUAUCAGCUGUAAGUUACUCAACUUUCCUAAAAGAUACUGAGCUUUACACGAUUAAAGAUGGUGAAGCUAUCGUAUUAUCGAGUAUUCCUUUUAAUGCAAAUUGGUUAAAUCAACAAUAUGCUGAAAUUAUCCAAGCAAUCUUAUUUGAUGUUGUAGGCUAUGAAGUUAAACCUCACUUUAUUACUACUGAAGAAUUAGCAAAUUAUAGUAAUAAUGAAACUGCUACUCCAAAAGAAACAACAAAACCUUCUACUGAAACAACUGAGGAUAAUCAUGUGCUUGGUAGAGAGCAAUUCAAUGCCCAUAACACAUUUGACACUUUUGUAAUCGGACCCGGUAACCGCUUUCCACAUGCAGCGAGUUUAGCUGUGGCCGAAGCACCAGCCAAAGCGUACAAUCCAUUAUUUAUCUAUGGAGGUGUUGGUUUAGGAAAAACCCAUUUAAUGCAUGCCAUUGGUCAUCAUGUUUUAGAUAAUAAUCCAGAUGCCAAAGUGAUUUACACAUCAAGUGAAAAAUUCACAAAUGAAUUUAUUAAAUCAAUUCGUGAUAACGAAGGUGAAGCUUUCAGAGAAAGAUAUCGUAAUAUCGACGUCUUAUUAAUCGAUGAUAUUCAGUUCAUACAAAACAAGGUACAAACACAAGAAGAAUUUUUCUAUACUUUUAAUGAAUUGCAUCAGAAUAACAAGCAAAUAGUUAUUUCGAGUGAUCGACCACCAAAGGAAAUUGCACAAUUAGAAGACCGAUUACGUUCACGCUUUGAAUGGGGGCUAAUUGUUGAUAUUACGCCACCAGAUUAUGAAACUCGAAUGGCAAUUUUGCAGAAGAAAAUUGAAGAAGAAAAAUUAGAUAUUCCACCAGAAGCUUUAAAUUAUAUAGCAAAUCAAAUUCAAUCUAAUAUUCGUGAAUUAGAAGGUGCAUUAACACGUUUACUUGCAUAUUCACAAUUAUUAGGAAAACCAAUUACAACUGAAUUAACUGCUGAAGCUUUAAAAGAUAUCAUUCAAGCACCAAAAUCUAAAAAGAUUACCAUCCAAGAUAUUCAAAAAAUUGUAGGCCAGUACUAUAAUGUUAGAAUUGAAGAUUUCAGUGCAAAAAAACGUACAAAGUCAAUUGCAUAUCCGCGUCAAAUAGCUAUGUACUUGUCUAGAGAGCUUACAGAUUUCUCAUUACCUAAAAUUGGUGAAGAAUUUGGUGGGCGUGAUCAUACGACCGUCAUUCAUGCUCAUGAAAAAAUAUCUAAAGAUUUAAAAGAAGAUCCUAUUUUUAAACAAGAAGUAGAGAAUCUUGAAAAAGAAAUAAGAAAUGUAUAA'
```


## Part 2

### Classical mechanics

The motion of a two-particle system can be decoupled into the motion of a the center-of-mass particle with an effective mass equal to the total mass of the system,


$$ {\bf r}_{cm} = \frac{m_1 {\textbf r}_1 + m_2 {\textbf r}_2}{m_1 + m_2} \qquad m_{tot} = m_1 + m_2$$
$$ $$

and the internal motion of a particle with a reduced mass

$$ {\bf r}_{red} = {\textbf r}_1 -  {\textbf r}_2 \qquad m_{red} = \frac{m_1m_2}{m_1 + m_2}$$



The bold-face notation ${\textbf r}$ refers to a point in space with $(x, y, z)$ coordinates and the scalar $m$ to the mass of a particle


### A Particle class

Construct a class of a particle with  $x, y, z$-coordinates and mass $m$. Let `x`, `y`, `z` and `m` be instance attributes of the class.

<!-- grade_id: cell-acdbbabfcf9a51c6, points: 0, grade: false, solution: true -->
```python
class Particle:
    """
    Class for particle with position and mass
    
    """
    # YOUR CODE HERE
    ################

```

<!-- grade_id: cell-0ee3ce1bc892e5dd, points: 1, grade: true, solution: false -->
```python
p = Particle(1, 2, 3, 4)
assert p.x == 1
assert p.y == 2
assert p.z == 3
assert p.m == 4
```

Define a `cm` method for a center-of-mass such that

<!-- grade_id: cell-6c4b4d28759ad052, points: 1, grade: true, solution: false -->
```python
p1 = Particle(0, 0, 0, 1)
p2 = Particle(0, 0, 1, 1)
r_cm = p1.cm(p2)
assert r_cm.x == 0
assert r_cm.y == 0
assert r_cm.z == 0.5
assert r_cm.m == 2
```

Define a reduced particle which satisfies

<!-- grade_id: cell-e6a55631ef9a1acc, points: 1, grade: true, solution: false -->
```python
p1 = Particle(0, 0, 0, 1)
p2 = Particle(0, 0, 1, 1)
r_red = p1.red(p2)

assert r_red.x == 0, f'{r_red.x=} != 0' 
assert r_red.y == 0, f'{r_red.y=} != 0' 
assert r_red.z == -1, f'{r_red.z=} != -1' 
assert r_red.m == 0.5
```

### Subclass: particle with charge

In this newclass the particle will have an additional attribute: charge, and the electrostatic interaction
between two charges is given by

$$ V = \frac{q_1 q_2}{|{\textbf r_1} - {\textbf r_2}|} $$

i.e. the product of the charges divided by the distance

Have a method `coulomb` in this new subclass that calculates this energy

<!-- grade_id: cell-5ce0017471fdb449, points: 0, grade: false, solution: true -->
```python
class ChargedParticle(Particle):
    """
    Class for particle with charge
    
    """
    # YOUR CODE HERE
    ################
```

### Test case: hydrogen atom

The electrostatic interaction between a proton and electron in a hydrogen atom

<!-- grade_id: cell-6186b66bd4288312, points: 1, grade: true, solution: false -->
```python
q1 = ChargedParticle(0,0,0,1, 1)
q2 = ChargedParticle(0,0,1,1, -1)
energy = q1.coulomb(q2)
expected = -1
assert energy == expected, f'{energy=} != {expected=}'
```

## Part 3

Given a DNA string write a generator that yield the characters in groups of three (codons).

<!-- grade_id: cell-36f00109bdfeb041, points: 0, grade: false, solution: true -->
```python
def codons(dna):
    """
    Yield the codons of a DNA/RNA string
    
    >>> for c in codons('GCAGCC'))
    ...      print(c)
    GCA
    GCC
    """
    # YOUR CODE HERE
    ################
```

Verify that this is a generator

<!-- grade_id: cell-58a7dcd4a99e5786, points: 1, grade: true, solution: false -->
```python
# Verify type
class_name = codons('').__class__.__name__
assert class_name == 'generator', f'{class_name=} != generator'
```

Verify some values in a defined string

```python
RNA = 'AUGUCGGAAAAAGAAAUUUGGGAAAAAGUGCUUGAAAUUGCUCAAGAAAAAUUAUCAGCUGUAAGUUACUCAACUUUCCUAAAAGAUACUGAGCUUUACACGAUUAAAGAUGGUGAAGCUAUCGUAUUAUCGAGUAUUCCUUUUAAUGCAAAUUGGUUAAAUCAACAAUAUGCUGAAAUUAUCCAAGCAAUCUUAUUUGAUGUUGUAGGCUAUGAAGUUAAACCUCACUUUAUUACUACUGAAGAAUUAGCAAAUUAUAGUAAUAAUGAAACUGCUACUCCAAAAGAAACAACAAAACCUUCUACUGAAACAACUGAGGAUAAUCAUGUGCUUGGUAGAGAGCAAUUCAAUGCCCAUAACACAUUUGACACUUUUGUAAUCGGACCCGGUAACCGCUUUCCACAUGCAGCGAGUUUAGCUGUGGCCGAAGCACCAGCCAAAGCGUACAAUCCAUUAUUUAUCUAUGGAGGUGUUGGUUUAGGAAAAACCCAUUUAAUGCAUGCCAUUGGUCAUCAUGUUUUAGAUAAUAAUCCAGAUGCCAAAGUGAUUUACACAUCAAGUGAAAAAUUCACAAAUGAAUUUAUUAAAUCAAUUCGUGAUAACGAAGGUGAAGCUUUCAGAGAAAGAUAUCGUAAUAUCGACGUCUUAUUAAUCGAUGAUAUUCAGUUCAUACAAAACAAGGUACAAACACAAGAAGAAUUUUUCUAUACUUUUAAUGAAUUGCAUCAGAAUAACAAGCAAAUAGUUAUUUCGAGUGAUCGACCACCAAAGGAAAUUGCACAAUUAGAAGACCGAUUACGUUCACGCUUUGAAUGGGGGCUAAUUGUUGAUAUUACGCCACCAGAUUAUGAAACUCGAAUGGCAAUUUUGCAGAAGAAAAUUGAAGAAGAAAAAUUAGAUAUUCCACCAGAAGCUUUAAAUUAUAUAGCAAAUCAAAUUCAAUCUAAUAUUCGUGAAUUAGAAGGUGCAUUAACACGUUUACUUGCAUAUUCACAAUUAUUAGGAAAACCAAUUACAACUGAAUUAACUGCUGAAGCUUUAAAAGAUAUCAUUCAAGCACCAAAAUCUAAAAAGAUUACCAUCCAAGAUAUUCAAAAAAUUGUAGGCCAGUACUAUAAUGUUAGAAUUGAAGAUUUCAGUGCAAAAAAACGUACAAAGUCAAUUGCAUAUCCGCGUCAAAUAGCUAUGUACUUGUCUAGAGAGCUUACAGAUUUCUCAUUACCUAAAAUUGGUGAAGAAUUUGGUGGGCGUGAUCAUACGACCGUCAUUCAUGCUCAUGAAAAAAUAUCUAAAGAUUUAAAAGAAGAUCCUAUUUUUAAACAAGAAGUAGAGAAUCUUGAAAAAGAAAUAAGAAAUGUAUAA'
```

<!-- grade_id: cell-ed847302a879f082, points: 1, grade: true, solution: false -->
```python
assert list(codons(RNA))[0] == 'AUG'
assert list(codons(RNA))[-1] == 'UAA'
```

### Count

Look up the `Counter` class in the `collections` module. It has a method `most_common`. Write a wrapper to that function so you get the most common codon and its count in the string

    >>> most_common_codon('AAAGGGAAA')
    ('AAA', 2)

<!-- grade_id: cell-425b86d1aa31e752, points: 0, grade: false, solution: true -->
```python
from collections import Counter

def most_common_codon(dna):
    """
    Returns the most common codon with its count
    >>> most_common_codon('AAAGGGAAA')
    ('AAA', 2)
    """
    # YOUR CODE HERE
    ################

```

<!-- grade_id: cell-07d22f7dfa7badd6, points: 1, grade: true, solution: false -->
```python
assert most_common_codon('AAAGGGAAA') == ('AAA', 2)
assert most_common_codon(RNA) == ('GAA', 42)
```

### Search

Write a function to search for a given substring and return the number (counting codons) of the first match according to

    >>> find_codon('AAA'  'AAAGGG')
    0
    >>> find_codon('GGG', 'AAAGGG')
    1

<!-- grade_id: cell-0f04b19956a83a07, points: 0, grade: false, solution: true -->
```python
def find_codon(codon, dna):
    """
    >>> find_codon('AAA'  'AAAGGG')
    0
    >>> find_codon('GGG', 'AAAGGG')
    1
    """
    # YOUR CODE HERE
    ################
```

<!-- grade_id: cell-532880ad9d9e4fe1, points: 1, grade: true, solution: false -->
```python
assert find_codon('AUG', RNA) == 0
assert find_codon('GAA', RNA) == 2
```