Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
PERSON_NUMBER = ""
```

---

# Tenta 2023-06-09 8:00-12:00



## Part 1

Here we use data from the git lab (git-2023-04-19). Use the current version of `booklist.txt` in what follows.

* The data in the file is close to a csv file but the formatting is not consistent all over: separation between author and title is often a comma, but somtimes by `-`, sometimes the word `by` and similar. Start by writing a script that cleans the data and saves the book in a pure csv file `booklist.csv`

* Let the new csv-file have semi-colon as separator, to allow for comma in title. In this sample we accept the simplication that there are no books with more than two authors.

(4p)

```python
!git clone git@gits-15.sys.kth.se:bb1000/git-2023-04-19
```

<!-- grade_id: cell-9a7763cb656ee883, points: 0, grade: false, solution: true -->
```python
# transform `booklist.txt` to `booklist.csv`
from pathlib import Path

source = Path('git-2023-04-19') / 'booklist.txt'
target = 'booklist.csv'

with open(target, 'w') as csvfile:
    csvfile.write('title;author\n')
    for line in open(source):
        # YOUR CODE HERE
        ################
```

### Verification

<!-- grade_id: cell-d9970f4369104a3d, points: 1, grade: true, solution: false -->
```python
# Verify the transformation of `by`:
# Pro Git by Scott Chacon and Ben Straub

books = open('booklist.csv').read()
assert "Pro Git;Scott Chacon and Ben Straub\n" in books
```

<!-- grade_id: cell-926e2f285977660e, points: 1, grade: true, solution: false -->
```python
# Verify the transformation of `-`
# Allegro Pastell - Leif Randt

assert "Allegro Pastell;Leif Randt\n" in books
```

<!-- grade_id: cell-098f595ec2ff62e8, points: 1, grade: true, solution: false -->
```python
# Verify that lacking author gets an `Unknown entry`
assert "Formel i hallfastighetslara;Unknown\n" in books
```

<!-- grade_id: cell-c259d79d4907b051, points: 1, grade: true, solution: false -->
```python
# Verify only one separator `;` per line in the transformed file
entries_per_line = [line.count(';') for line in open('booklist.csv')]
assert set(entries_per_line) == {1}, entries_per_line
```


## Part 2

* Use the generated csv file to create a dictionary that maps a title to an author. Implement this in a function `csv_to_dict`


<!-- grade_id: cell-5f2ab8ed945c9bf7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
################
```

<!-- grade_id: cell-d9c9c1c0fa57a858, points: 1, grade: true, solution: false -->
```python
books = csv_to_dict('booklist.csv', sep=';')
assert type(books) == dict
assert books['1984'] == 'George Orwell'
```

* Dictinaries preserve the order in which they were created. Make sure that the data is sorted by title so that there first entry in the dictionary is the *1984* novel

<!-- grade_id: cell-da3a9a0a9df1c647, points: 1, grade: true, solution: false -->
```python
assert list(books.items())[0] == ('1984', 'George Orwell')
```

* Reimplement these concepts into a class such that

<!-- grade_id: cell-485dea1be87db90b, points: 0, grade: false, solution: true -->
```python
class Books:
    # YOUR CODE HERE
    ################
```

<!-- grade_id: cell-41e269deeedf1b19, points: 1, grade: true, solution: false -->
```python
books = Books('booklist.csv')

assert list(books.to_dict().items())[0] == ('1984', 'George Orwell')
```

<!-- grade_id: cell-249a6af80039fa30, points: 1, grade: true, solution: false -->
```python
assert books['1984'] == 'George Orwell', books['1984']
```

## Part 3

We build on the previous class by adding search functions. Let the instance method `find_title` such that

~~~
>>> for title in books.find_title('python'):
...     print(title)
Automate the Boring Stuff with Python
Python Cookbook
Python For Dummies
~~~

i.e. let the matching ignore case. To anticipate very large files this should result in a generator

<!-- grade_id: cell-61326ce87f60dbc4, points: 0, grade: false, solution: true -->
```python
def find_title(self, search):
    # YOUR CODE HERE
    ################
    
Books.find_title = find_title
books = Books('booklist.csv')
```

<!-- grade_id: cell-696354e140f231e4, points: 1, grade: true, solution: false -->
```python
# is it a generator?
python_books = books.find_title('python')
assert iter(python_books) == python_books
```

<!-- grade_id: cell-e4eb983cc2c733a9, points: 1, grade: true, solution: false -->
```python
python_books = books.find_title('python')
assert list(python_books) ==[
    'Automate the Boring Stuff with Python',
    'Python Cookbook',
    'Python For Dummies'
]
```

* Build on the previous example to search in the author field and return the titles for case-insensitive author matches

<!-- grade_id: cell-122170bfa7ee520d, points: 0, grade: false, solution: true -->
```python
def find_author_titles(self, search):
    # YOUR CODE HERE
    ################
    
Books.find_author_titles = find_author_titles
books = Books('booklist.csv')
```

<!-- grade_id: cell-e76a631471721456, points: 1, grade: true, solution: false -->
```python
assert list(books.find_author_titles('Strindberg')) == ['En dåres försvarstal', 'Roda rummet']
```

* Define a decorator that takes a function returning a dictionary to instead return a pandas Series object. We apply this to the to_dict method of our class as below

<!-- grade_id: cell-2482495dad6a826d, points: 0, grade: false, solution: true -->
```python
import pandas as pd
def dict_to_series(f):
    # YOUR CODE HERE
    ################
```

```python
Books.to_series = dict_to_series(Books.to_dict)
books = Books('booklist.csv')
```

<!-- grade_id: cell-e019b7ccb9c03c5c, points: 1, grade: true, solution: false -->
```python
assert type(books.to_dict()) == dict
assert type(books.to_series()) == pd.Series
assert books.to_series()['1984'] == 'George Orwell'
```