Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
PERSONNUMBER = ""
```

---

# Exam in BB1000, 2021-08-20

## Grading

There are three levels in order of increasing difficulty

* Grade D for 3 of 4 points at level 1 and 1 point at level 2
* Grade C for 3 of 4 points at levels 1 and 2
* Grade B as C plus 1 point at level 3
* Grade A for 3 of 4 points at all levels

# Requirements

For the whole exam you should be able to import the following libraries

```python
import pandas
```

# Level 1


**1. The file `market-price.csv` contains historical daily bitcoin data for the recent year. Write a function that extracts the value in dollars for a given date**

This can be a simple implementation looping over the lines of the file processing the line
with the matching date. Not that the return value is expected to be of type float

~~~
>>> get_usd('2021-04-01', 'market-price.csv')
58735.25
~~~

The file has the structure

    2020-08-20 00:00:00,11734.0
    2020-08-21 00:00:00,11865.82
    2020-08-22 00:00:00,11522.8
    ...

<!-- grade_id: cell-f822b8fd33ae224e, points: 0, grade: false, solution: true -->
```python
def get_usd(date_string, filename):
    """
    Returns exchange rate in file for given date
    """
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-6d536596d441b56e, points: 1, grade: true, solution: false -->
```python
result = get_usd('2021-04-01', 'market-price.csv')
expected = 58735.25
assert result == expected, f'{result} != {expected}'

assert get_usd('2021-08-19', 'market-price.csv') == 44777.86
```

**2. Read tha data values to a dictionary**

Let the date format be a string in isoformat YYYY-MM-DD, which is the dictionary key and the exchange rate the
dictionary value


<!-- grade_id: cell-fff1cc445821fd2a, points: 0, grade: false, solution: true -->
```python
def get_usd_as_dict(filename):
    """
    Returns a dictionary with dates as keys and exchange rate as values
    """
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-46b7195daa3616fd, points: 1, grade: true, solution: false -->
```python
usd = get_usd_as_dict('market-price.csv')
assert usd['2021-04-01'] == 58735.25
```

**3. At which date was the exchange rate highest?**

Write a function which takes a dictionary of the form of assignment 2 and returns the date with the highest rate.

Hint: look up the documentation for `sorted`.



<!-- grade_id: cell-d6f7f3efe0d9d995, points: 0, grade: false, solution: true -->
```python
def get_index_max_value(dct):
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-f3e3cea29770d038, points: 0, grade: true, solution: false -->
```python
assert get_index_max_value(usd) == '2021-04-14'
assert usd['2021-04-14'] == 63554.44
```

**4. Dollar cost averaging**

Dollar cost averaging (DCA) means to save a fixed amount in dollars on a regular basis. An asset with large variation in value over time is high risk saving, but DCA is beleived to reduce risk in a longer perspective. Assume that you save 1 dollar per day converting it to btc without any exchange fees. What is the total btc accumulated  during this year.


<!-- grade_id: cell-cb69b552e147b3dc, points: 0, grade: false, solution: true -->
```python
def dca(dct):
    """
    Return total btc acuumulated from $1/day
    """
    # YOUR CODE HERE
    
```

<!-- grade_id: cell-6f1176f4ef0dbb1c, points: 1, grade: true, solution: false -->
```python
assert dca({'2008-01-01': 10}) == .1
assert dca({'2008-01-01': 10, '2008-01-02': 40}) == .125
assert round(dca(usd), 8) == 0.01529260  # values is a dict from assignment 2
```

*Observation*:
the total invested and  value at the end of the year

```python
print('The total invested:', len(usd))
print('Value after one year', round(dca(usd) * list(usd.values())[-1], 2))
```

# Level 2

**5. Use of pandas**

Read in the data in `market-price.csv` as a pandas DataFrame. Initially it may look ill-formatted, if the data file does not have headers if its own

          2020-08-20 00:00:00 	11734.0
    0 	2020-08-21 00:00:00 	11865.82
    1 	2020-08-22 00:00:00 	11522.80
    2 	2020-08-23 00:00:00 	11683.44
    ... 
  

<!-- grade_id: cell-5bed954dde7d9e1e, points: 0, grade: false, solution: true -->
```python
import pandas
data = (
    # YOUR CODE HERE
    
)
data.head()
```

<!-- grade_id: cell-86051b003641f730, points: 1, grade: true, solution: false -->
```python
assert isinstance(data, pandas.DataFrame)
```

**6. Find options in 5. so that column headers are 'date' and 'USD'**

          date 	                USD
    0 	2020-08-20 00:00:00 	11734.00
    1 	2020-08-21 00:00:00 	11865.82
    2 	2020-08-22 00:00:00 	11522.80
    3 	2020-08-23 00:00:00 	11683.44
    4 	2020-08-24 00:00:00 	11653.02

<!-- grade_id: cell-c85f802cde279ff8, points: 0, grade: false, solution: true -->
```python
data = (
    # YOUR CODE HERE
    
)
data.head()
```

<!-- grade_id: cell-3cb30633bae04d4f, points: 1, grade: true, solution: false -->
```python
assert all(data.columns == ['date', 'USD'])
```

**7. How du you make the date column a dataframe index?**

                             USD
    date 	
    2020-08-20 00:00:00 	11734.00
    2020-08-21 00:00:00 	11865.82
    2020-08-22 00:00:00 	11522.80
    ...
    

<!-- grade_id: cell-96206ee5afd04747, points: 0, grade: false, solution: true -->
```python
data = (
    # YOUR CODE HERE
    
)
data.head()
```

<!-- grade_id: cell-253d7bd6c66944ec, points: 1, grade: true, solution: false -->
```python
assert data.columns == ['USD']
assert data.index.name == 'date'
data.dtypes
```

**8. convert the expression to a Series object**

    date
    2020-08-20    11734.00
    2020-08-21    11865.82
    2020-08-22    11522.80
    ...
    Name: USD, dtype: float64

<!-- grade_id: cell-cded83a7864ffa49, points: 0, grade: false, solution: true -->
```python
data = (
    # YOUR CODE HERE
    
)
data.head()
```

<!-- grade_id: cell-969b94b2fb1f7424, points: 1, grade: true, solution: false -->
```python
assert isinstance(data, pandas.Series)
assert data.name == 'USD'
assert data.index.name == 'date'
assert data.dtype == 'float'
```

*Observation*: date-like entries can be processed with the option parse_dates=[0], so that Python can internally use date-types. See e.g. the difference it makes for plotting. Use this in next assignment

```python
data.plot()
```

# Level 3

The goal here is to obtain a decorator that returns values from `market-price.csv` in SEK insted of USD given the exchange rate for a given date in the file `exchange-rates.csv`



<!-- grade_id: cell-6869503264cfa8fb, points: 0, grade: false, solution: true -->
```python
exchange_rates = (
    # YOUR CODE HERE
    
)
exchange_rates.head()
```

**9.** Verify that the `Date` column has the pandas datatype `Timestamp`

<!-- grade_id: cell-55cff1d63b590d76, points: 1, grade: true, solution: false -->
```python
assert isinstance(exchange_rates.iloc[0]['Date'], pandas.Timestamp)
```

**10.** Create a column with dates in string iso format. The result of the following cell is us2sek as a pandas.Series object

*Hint: work with the `apply` method of a series and a function that uses datetime.isoformat*

<!-- grade_id: cell-6747c8c5721d3d33, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

exchange_rates.set_index('isodate', inplace=True)
us2sek = exchange_rates['SEK']
us2sek
```

<!-- grade_id: cell-ea674f60101cc06e, points: 1, grade: true, solution: false -->
```python
assert isinstance(us2sek, pandas.Series)
assert us2sek['2020-08-18'] == 8.6679
```

**11.** Complete a decorator so that `get_sek` returns the btc value in SEK instead of dollars

<!-- grade_id: cell-d56d07f0f8902c20, points: 0, grade: false, solution: true -->
```python

def convert(f):
    # YOUR CODE HERE
    


get_sek = convert(get_usd)
        
```

<!-- grade_id: cell-0e185dc400c08618, points: 1, grade: true, solution: false -->
```python
assert get_sek('2020-08-20', 'market-price.csv').round(2) == 102541.08
```

<!-- grade_id: cell-d40ea7172d8e4cdb, points: 1, grade: true, solution: false -->
```python
assert get_sek('2021-08-18', 'market-price.csv').round(2) == 389366.78
```