Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
PERSON_NUMBER = ""
```

---

# Tenta BB1000 2022-06-01

## Grading

There are three levels in order of increasing difficulty

    Grade E for 75% at level 1
    Grade D as E plus 25% at level 2
    Grade C for 75% at levels 1 and 2
    Grade B as C 25% at level 3
    Grade A for 75% at all levels

For the exam you should be able to import the following external libraries: requests


## Part 1

**1**. There is a command-line interface to the Google calendar that allows you to view and add events. If you have it installed (NOT REQUIRED FOR THIS ASSIGNMENT) help for adding events can be read with

    $ gcalcli add -h
    usage: gcalcli add [-h] [--details {calendar,location,length,reminders,description,url,attendees,email,attachments,end,all}] [--reminder REMINDERS] [--default-reminders] [--color EVENT_COLOR]
                       [--title TITLE] [--who WHO] [--where WHERE] [--when WHEN] [--duration DURATION] [--description DESCRIPTION] [--allday] [--noprompt]

Suppose you would like to add things to your calendar from the command line but you find that the interface is too inconvenient. You decide to generate the command from Python where some options have reasonable defaults. You decide to setup the  command-line options from a Python dictionary, such that e.g.

    {"when": "now"}
    
makes a contribution to the command with

    '--when "now"'


You find that it makes most sense to always have the dictionary values in quotation marks to allow for e.g. multi-word titles with spaces. 

Write a `generate_command` function which accepts  keyword arguments into a dictionary. As a first testcase
verify that

~~~
>>> generate_command(title='exam', when='8:00')
'gcalcli add --title "exam" --when "8:00"'
~~~

How do you continue with the implementation?

<!-- grade_id: cell-09994c5a863df0e2, points: 0, grade: false, solution: true -->
```python
def generate_command(**keywords):
    """
    Generate bash command
    
    >>> generate_cmd(when="12:00", what="do this")
    gcalcli add --when "12:00" --what "do this"
    """
    # YOUR CODE HERE
    ################
    

# Example
generate_command(title='exam', when='8:00')
```

<!-- grade_id: cell-bf4c38be3cee98ee, points: 1, grade: true, solution: false -->
```python
cmd = generate_command(title='exam', when='8:00') 
expected = 'gcalcli add --title "exam" --when "8:00"'
assert cmd == expected, f'\nActual:   {cmd}\nExpected: {expected}'
```

**2**. When the calendar event does not explicitly include time the command line should contain the '--allday' flag as well. Assume that this holds whenever the colon character `":"` is not part of the `when=` keyword argument

~~~
>>> generate_command(title='study more', when='tomorrow')
'gcalcli add --title "study more" --when "tomorrow" --allday'
~~~

<!-- grade_id: cell-cc36ecc57ca42fd1, points: 1, grade: true, solution: false -->
```python
cmd = generate_command(title='study more', when='tomorrow') 
expected = 'gcalcli add --title "study more" --when "tomorrow" --allday'
assert cmd == expected, f'\nActual:   {cmd}\nExpected: {expected}'
```

**3**.
Some of the keywords can be predifined in a configuration file. Set a default calendar with the json file

```python
%%file config.json
{
    "calendar":"user@gmail.com"
}
```

 and read it into a Python dictionary `config` such that
 
     >>> get_config('config.json')
     {'calendar': 'user@gmail.com'}

<!-- grade_id: cell-13cb5f8b86ce5e89, points: 0, grade: false, solution: true -->
```python
import json
def get_config():
    # YOUR CODE HERE
    ################
    
get_config()
```

Verify now that the following holds

    >>> get_command(**get_config())
    'gcalcli add --calendar "user@gmail.com"'
    

<!-- grade_id: cell-473a850678f55468, points: 1, grade: true, solution: false -->
```python
cmd = generate_command(**get_config())
expected = 'gcalcli add --calendar "user@gmail.com"'
assert cmd == expected, f'\nActual:   {cmd}\nExpected: {expected}'
```

**4** If the dictionary value is the None object let that translate to an empty string

    >>> generate_command(description=None)
    'gcalcli add --description ""'

<!-- grade_id: cell-342dbd0eca8b37a9, points: 1, grade: true, solution: false -->
```python
cmd = generate_command(description=None)
expected = 'gcalcli add --description ""'
assert cmd == expected, f'\nActual:   {cmd}\nExpected: {expected}'
```

## Part 2

A User class is used to save username and password. The instance attributes are `username` and `_password`. The stored password should be encrypted and for that we use the helper function `encrypt`. The c`users` class attribute is a dictionary that contains all created users

```python
import hashlib
def encrypt(password):
    return hashlib.md5(password.encode()).hexdigest()

encrypt('querty')
```

<!-- grade_id: cell-92758d01db883240, points: 0, grade: false, solution: true -->
```python
class User:
    """
    instance attributes
    
    username: str
    _password: str
    
    class attributes
    
    users: dict
    
    """
    # YOUR CODE HERE
    ################
```

**5** Write an initializer such that the following hold

<!-- grade_id: cell-3480ef7510acc50e, points: 1, grade: true, solution: false -->
```python
adam = User('Adam', 'qwerty')
assert adam.username == 'Adam'
assert adam._password != 'querty'
```

**6** Write a string representation function for the User class such that

<!-- grade_id: cell-3f6c49612826ea93, points: 1, grade: true, solution: false -->
```python
assert repr(adam) == 'User("Adam")'
```

**7** Verify that a supplied password in a login situation matches the stored encrypted password

<!-- grade_id: cell-a2a2cc05fa5905a1, points: 1, grade: true, solution: false -->
```python
assert adam.verify('qwerty')
assert not adam.verify('qwertu')
```

**8** include a method that lists the other users

```python
eve = User('Eve', 'asdfgh')
kain = User('Kain', 'omg')
abel = User('Abel', 'foo')
```

<!-- grade_id: cell-8eb8f37aad90db25, points: 1, grade: true, solution: false -->
```python
actual = adam.friends()
expected = ['Eve', 'Kain', 'Abel']
assert actual == expected, f'\nActual:   {actual}\nExpected: {expected}'

```

## Part 3


KTH provides an web application programming interface (API) for the schedule database. 

The following command returns a Python dictionary with some nested data for the time period Jun-Aug 22

    >>> requests.get('https://kth.se/api/schema/v2/course/bb1000?startTime=2022-06-01&endTime=2022-08-31').json()
    {'url': 'https://www.kth.se/social/course/BB1000/calendar/',
     'entries': [{'url': 'https://www.kth.se/social/course/BB1000/subgroup/vt-2022-811/event/556204/',
         'start': '2022-06-01 08:00:00',
         'end': '2022-06-01 12:00:00',
         'title': 'Tentamen',
         'type': 'Ten',
         'type_name': {'sv': 'Tentamen', 'en': 'Examination'},
         'group': '',
         'locations': [{'name': 'FB42',
             'url': 'https://www.kth.se/places/room/id/d0adc673-3b21-4590-8e07-4995c3c449df'},
            {'name': 'FB52',
             'url': 'https://www.kth.se/places/room/id/8af44ca1-c331-43fe-b1bf-437047c7afa4'},
            {'name': 'FD41',
             'url': 'https://www.kth.se/places/room/id/fc59066e-51b3-4980-ba39-1a9d508ad8a8'}]},
        ...

(which you can verify)

```python
import requests
KTH_SCHEMA = 'https://kth.se/api/schema/v2/course/bb1000?startTime=2022-06-01&endTime=2022-08-31'
requests.get(KTH_SCHEMA).json()
```

Complete the function below to be a generator that returns dictionaries with information from kth but in a form useful for the Google calendar function in part 1.

In particular we need to extract

    * start time of an event (for the --when option)
    * location (for the --where option)
    * the event title (for the --what option)
    
For the case that we have more than one location it is enough to consider the first/

<!-- grade_id: cell-8cd12e4734ffa3b4, points: 0, grade: false, solution: true -->
```python
KTH_URL = 'https://kth.se/api/schema/v2/course/'
def kth2gcal(course, start_time='2022-06-01', end_time='2022-08-31'):
    # YOUR CODE HERE
    ################
```

**9**. This verifies that `kth2gcal is a generator

<!-- grade_id: cell-7fa089cac31daa9a, points: 1, grade: true, solution: false -->
```python
import typing
assert isinstance(kth2gcal, typing.Callable)  # is a function
assert isinstance(kth2gcal(''), typing.Generator)  # is a generator
```

**10.** Then we loop ever all events of the period the last one is the August reexam

<!-- grade_id: cell-02c75c088cda8bc2, points: 1, grade: true, solution: false -->
```python
for event in kth2gcal('BB1000', start_time='2022-06-01', end_time='2022-08-31'):
    pass

print(event)

expected = dict(where='FB41', what='Omtenta', when='2022-08-19 08:00:00')
assert event == expected, f'\nActual: {event}\nExpected {expected}'
```

**11**. Verify that this fits in with the gcalcli command

<!-- grade_id: cell-be2b67c958e0db9b, points: 1, grade: true, solution: false -->
```python
cmd = generate_command(**event)
print(cmd)
assert '--where "FB41"' in cmd
assert '--what "Omtenta"' in cmd
```

**12**. A function `search_lab_date` finds the  occurances of a lab session where a given computer room has been scheduled. 

Write an expression involving kth2gcal which returns a list of such dates which is returned by the function

<!-- grade_id: cell-a15015dd5c54da2d, points: 0, grade: false, solution: true -->
```python
def search_lab_dates(class_room):
# YOUR CODE HERE
################
```

<!-- grade_id: cell-435eec3d3f7af324, points: 1, grade: true, solution: false -->
```python
actual = search_lab_dates('Gul')[0]  # The first lab in yellow
expected = '2022-05-20 13:00:00'
assert actual == expected, f'\nActual   {actual}\nExpected {expected}'
```