# Tenta 2020-08-14

## Part 1

### 1 Reading email data

Consider the file written below as staff.csv. It is a CSV-file (comma-separated values) with random personal data: 

    1,Conn,De Gregoli,cdegregoli0@epa.gov,"Kunze, Crona and Bergnaum"
    2,Anette,Stutt,astutt1@oracle.com,Kuhlman LLC
    3,Eberto,Herkess,eherkess2@google.com.au,Russel Inc
    4,Montgomery,Dublin,mdublin3@dion.ne.jp,Gulgowski Group
    5,Jenilee,Readwing,jreadwing4@craigslist.org,Hansen and Sons
    6,Brit,Scattergood,bscattergood5@wiley.com,Greenholt and Sons
    7,Abby,Minkin,aminkin6@purevolume.com,"Krajcik, Reichert and MacGyver"
    8,Aleda,Bradburne,abradburne7@quantcast.com,"Hintz, Lang and Koepp"
    9,Abramo,Shoemark,ashoemark8@yahoo.co.jp,"O'Keefe, Morissette and Erdman"
    10,Inness,Gibbins,igibbins9@geocities.com,Bartell Inc

Create a function `extract_emails` that takes a filename as input and returns a list of strings containing the email-addresses.


<!-- grade_id: cell-ecf3707429c73ab6, points: 0, grade: false, solution: true -->
```python
def extract_emails(filename):
    # YOUR CODE HERE
    ...
```

We first verify the behaviour of a single-line file, then with the full staff file

```python
%%file 'first.csv'
1,Conn,De Gregoli,cdegregoli0@epa.gov,"Kunze, Crona and Bergnaum"
```

Output:
```
Overwriting first.csv
```

```python
%%file 'staff.csv'
1,Conn,De Gregoli,cdegregoli0@epa.gov,"Kunze, Crona and Bergnaum"
2,Anette,Stutt,astutt1@oracle.com,Kuhlman LLC
3,Eberto,Herkess,eherkess2@google.com.au,Russel Inc
4,Montgomery,Dublin,mdublin3@dion.ne.jp,Gulgowski Group
5,Jenilee,Readwing,jreadwing4@craigslist.org,Hansen and Sons
6,Brit,Scattergood,bscattergood5@wiley.com,Greenholt and Sons
7,Abby,Minkin,aminkin6@purevolume.com,"Krajcik, Reichert and MacGyver"
8,Aleda,Bradburne,abradburne7@quantcast.com,"Hintz, Lang and Koepp"
9,Abramo,Shoemark,ashoemark8@yahoo.co.jp,"O'Keefe, Morissette and Erdman"
10,Inness,Gibbins,igibbins9@geocities.com,Bartell Inc
```

Output:
```
Overwriting staff.csv
```

<!-- grade_id: cell-645281897a6430d7, points: 1, grade: true, solution: false -->
```python
from solutions import extract_emails
assert extract_emails('first.csv') == ['cdegregoli0@epa.gov']
```

<!-- grade_id: cell-d5342fbdaa2cf7b4, points: 1, grade: true, solution: false -->
```python
assert extract_emails('staff.csv') == [
    'cdegregoli0@epa.gov',
     'astutt1@oracle.com',
     'eherkess2@google.com.au',
     'mdublin3@dion.ne.jp',
     'jreadwing4@craigslist.org',
     'bscattergood5@wiley.com',
     'aminkin6@purevolume.com',
     'abradburne7@quantcast.com',
     'ashoemark8@yahoo.co.jp',
     'igibbins9@geocities.com'
]
```

### 2 Pass emails to a mail program

Reformat the output `extract_emails` to have email addresses as a string separated by a semicolon

<!-- grade_id: cell-c07613855916b690, points: 0, grade: false, solution: true -->
```python
def sendmail_format(filename):
    # YOUR CODE HERE
    ...
```

<!-- grade_id: cell-3e0b0f5f82a66ea6, points: 1, grade: true, solution: false -->
```python
from solutions import sendmail_format
assert sendmail_format('first.csv') == 'cdegregoli0@epa.gov'
```

<!-- grade_id: cell-1613a36a3e523031, points: 1, grade: true, solution: false -->
```python
assert sendmail_format('staff.csv') == 'cdegregoli0@epa.gov; astutt1@oracle.com; eherkess2@google.com.au; mdublin3@dion.ne.jp; jreadwing4@craigslist.org; bscattergood5@wiley.com; aminkin6@purevolume.com; abradburne7@quantcast.com; ashoemark8@yahoo.co.jp; igibbins9@geocities.com'

```

```{admonition} Solution
:class: dropdown
~~~
def extract_emails(filename):                                                              
    with open(filename)as f:                                                               
        return [line.split(',')[3] for line in f]                                          
~~~
```

## Part 2

### 1. An employee class
A class initializer taker as input a line of the same format as in the staff file of part 1. Define instance attributes for the class such that the test cells are valid

<!-- grade_id: cell-5c2f8aaede02369a, points: 0, grade: false, solution: true -->
```python
class Employee:
    # YOUR CODE HERE
    ...
```

<!-- grade_id: cell-d93247ba92e31997, points: 1, grade: true, solution: false -->
```python
from solutions import Emp as Employee
employee = Employee('1,Conn,De Gregoli,cdegregoli0@epa.gov,"Kunze, Crona and Bergnaum"')
assert employee.first == 'Conn'
```

<!-- grade_id: cell-5c2f8aaede02369b, points: 1, grade: true, solution: false -->
```python
assert employee.last == 'De Gregoli'
```

<!-- grade_id: cell-7bce9cbd985f9902, points: 1, grade: true, solution: false -->
```python
assert employee.email == 'cdegregoli0@epa.gov'
```

```{admonition} Solution
:class: dropdown
~~~
class Employee:                                                                            
    def __init__(self, line):                                                              
        self.first, self.last, self.email = line.split(',')[1:4]                           
~~~
```

### 2. Text representation
Mail programs accept mail receiver in the format "first last <email>". Update the Employee class with a function such that
    
    str(employee) == 'Conn De Gregoli <cdegregoli>'

<!-- grade_id: cell-acbef7ff1fc24536, points: 1, grade: true, solution: false -->
```python
assert str(employee) == 'Conn De Gregoli <cdegregoli0@epa.gov>'
```

```{admonition} Solution
:class: dropdown
~~~
class Employee:                                                                            
    ...
    def __str__(self):                                                                     
        return f'{self.first} {self.last} <{self.email}>'                                  
~~~
```

## Part 3

### 1. Reusing Employee

Reuse the class definition Employee of part 2 (copy-paste your previous solution) and the staff file

<!-- grade_id: cell-593e7065a23e9534, points: 0, grade: false, solution: true -->
```python
# class Employee:
    # YOUR CODE HERE
```

```python
%%file 'staff.csv'
1,Conn,De Gregoli,cdegregoli0@epa.gov,"Kunze, Crona and Bergnaum"
2,Anette,Stutt,astutt1@oracle.com,Kuhlman LLC
3,Eberto,Herkess,eherkess2@google.com.au,Russel Inc
4,Montgomery,Dublin,mdublin3@dion.ne.jp,Gulgowski Group
5,Jenilee,Readwing,jreadwing4@craigslist.org,Hansen and Sons
6,Brit,Scattergood,bscattergood5@wiley.com,Greenholt and Sons
7,Abby,Minkin,aminkin6@purevolume.com,"Krajcik, Reichert and MacGyver"
8,Aleda,Bradburne,abradburne7@quantcast.com,"Hintz, Lang and Koepp"
9,Abramo,Shoemark,ashoemark8@yahoo.co.jp,"O'Keefe, Morissette and Erdman"
10,Inness,Gibbins,igibbins9@geocities.com,Bartell Inc
```

Output:
```
Overwriting staff.csv
```

Use the map function to generate a list of Employee objects

<!-- grade_id: cell-7c78b621c8b6307c, points: 0, grade: false, solution: true -->
```python
def read_list_of_employees(filename):
    # YOUR CODE HERE
    ...
```

<!-- grade_id: cell-48920e1b30b4a01e, points: 1, grade: true, solution: false -->
```python
## empty file
from solutions import read_list_of_employees
with open('empty.csv', 'w'):
    pass

assert read_list_of_employees('empty.csv') == []
```

<!-- grade_id: cell-2814fe462c032507, points: 1, grade: true, solution: false -->
```python
assert read_list_of_employees('staff.csv')[0].first == 'Conn'
assert read_list_of_employees('staff.csv')[1].first == 'Anette'
```

```{admonition} Solution
:class: dropdown
~~~
def read_list_of_employees(filename):
    return list(map(Employee, open(filename)))
~~~
```

### 2. filtered sendlist

Use the filter function to extract the employees with a japanese email address and generate the sendlist. 

Hint: you need a function that returns True of False depending on the address ends with '.jp' or not

<!-- grade_id: cell-cfe20e7052343436, points: 0, grade: false, solution: true -->
```python
def japan_sendlist(filename):
    employees = read_list_of_employees(filename)
    # YOUR CODE HERE
```

<!-- grade_id: cell-33d73c54a77e858b, points: 1, grade: true, solution: false -->
```python
from solutions import japan_sendlist
assert japan_sendlist('empty.csv') == ''
```

<!-- grade_id: cell-5db6565d8233d895, points: 1, grade: true, solution: false -->
```python
assert japan_sendlist('staff.csv') == 'Montgomery Dublin <mdublin3@dion.ne.jp>; Abramo Shoemark <ashoemark8@yahoo.co.jp>'
```

```{admonition} Solution
:class: dropdown
~~~
def japan_sendlist(filename):
    # Define a helper function for filtering japanese email addresses
    def is_jap(employee):
        return employee.email.endswith('.jp')
        
    filtered = filter(is_jap, employees)
    return '; '.join(map(str, filtered))
~~~
```